Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Evan"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work 11.2

## Part A.  Currency trading (Slightly simplified version of 15.3-6 from Cormen et al.)

Imagine that you wish to exchange one currency for another. You realize that instead of directly exchanging one currency for another, you might be better off making a series of trades through other currencies, winding up with the currency you want.

Suppose that you can trade n different currencies, numbered $1,2,… ,n$, where you start with currency 1 and wish to wind up with currency $n$. You are given, for each pair of currencies $i$ and $j$ , an exchange rate $r_{ij}$ , meaning that if you start with $d$ units of currency $i$ , you can trade for $dr_{ij}$ units of currency $j$. Note that the total number of trades allowed is limited to $n$.

Assuming there is no commission, propose an algorithm (described in prose) to solve this problem using either a memoization or bottom-up strategy. **Note that we will be working on the Python implementation in class.**

Click [here](https://drive.google.com/open?id=1L8Fjo1Xt8sltab-tz3m91eTuiE9LYWF7) for some example data.

<!-- grade_id: cell-3f8a8808c720baf2, points: 0, grade: true, solution: true -->
This problem could be solved by using memoization to record the subproblems of getting from currency 1 to n. We will have n-1 currecies that the first currency can be traded for. We could save the highest possible returns for a trade for all currencies on an nxn matrix. Each cell in the matrix would represent not just the direct exchange rates (as given) but the exchange rates for the optimal solution. This is longest weighted path directed network problem. In some cases, the most profitable path would be the direct path, and in others it would be an indirect path. The matrix would call on this memory rather than computing all the possible paths explicity to create an efficient solution.

## Part B - Money game (Solution)
Consider a row of n coins of values $v_1, v_2,...,v_n$, **where $n$ is even**. We play a game against an opponent by alternating turns. In each turn, a player selects either the first or last coin from the row, removes it from the row permanently, and receives the value of the coin. Determine the maximum possible amount of money we can definitely win if we move first.

For example, consider the game:

\$2, \$10, \$1, \$5

By moving first and playing optimally one can be guaranteed of \$15. The first move is to take \$5. This forces your opponent to take either \$2 or \$1, and then allows you to take \$10. Assume that the opposing player also plays optimally (i.e., minimizing the gain of the first-move player.)

As a hint, we will provide the recurrence for the solution to this problem. Suppose $c(i,j)$ is the maximum possible amount of money the first player can win for the row of coins of value $v_i, v){i+1},...,v_j$ (so to solve our problem, we need to compute $c(1,n)$. Then:

$$c(i,j)=max[v_i+min(c(i+2,j), c(i+1,j-1)), v_j+min(c(i,j-2), c(i+1,j-1))]$$

You are also encouraged to watch [this video](https://www.youtube.com/watch?v=KnP8_L13xW4&list=PLF_a-qBXTGFektoI6JUOTRL36JlvD04BR&index=5) if you need some more hints.

## Question 1. 
Complete the following function to solve the game using a bottom-up strategy, assuming that the opposing player also plays optimally. 


<!-- grade_id: cell-4014ae901b66304b, points: 0, grade: false, solution: true -->
```python
def bottom_up_coin_game(A):
    """
    Returns the maximum possible amount of money the first-move player can win,
    given an array of coin values. The function solves this using a bottom-up 
    approach.
    
    Inputs:
    - A: list of floats, values of the coins, of even length. 
    
    Outputs:
    - max_val: float, maximum possible amount of money that can be won by the first 
    player. max_val is None when the length of list A is odd.
    """
    # YOUR CODE HERE
    n = len(A)
    if n%2 != 0:
        return None
    elif n == 0:
        return 0
    sol = []
    
    for i in range(n):
        sol.append([None]*(i+1))
    
    for i in range(n):

    
        if i == 0:
            for j in range(n):
                sol[j][-1]=A[j]

        if i == 1:
            for j in range(1,n):
                sol[j][-2]=max(A[j],A[j-1])
        
        print(sol)
        if i >1:
            for j in range(i,n):
                sol[j][-(i+1)]=max(A[j]+min(sol[j-2][-i+1],sol[j-1][-i+1]),A[j-i]+min(sol[j][-i+1],sol[j-1][-i+1]))

            #         for i in range(2,n):
#             sol[i][-3]=max(A[i],A[i-2])
        
#         for i in range(3,n):
#             sol[i][-4]=max(A[i],A[i-3])
            
            
    #c(i,j) = max(A[i]+min(c(i+2,j),c(i+1,j-1)),A[j]+min(c(i,j-2),c(i+1,j-1)))    
    return sol[-1][0]
                                                                                  
                                                                                  
    raise NotImplementedError()
```

```python
a = [] 
b=[2, 10, 1, 5]
for i in range(4):
    a.append([None]*(i+1))
a[2][-2]=max(b[2],b[2-1])    
print(a)
```

Output:
```
[[None], [None, None], [None, 10, None], [None, None, None, None]]
```

```python
solution_matrix =[[None for i in range(4)] for j in range(4)] 
print(solution_matrix)
```

Output:
```
[[None, None, None, None], [None, None, None, None], [None, None, None, None], [None, None, None, None]]
```

```python
print(bottom_up_coin_game([2, 10, 1, 5]))
```

Output:
```
[[2], [None, 10], [None, None, 1], [None, None, None, 5]]
[[2], [10, 10], [None, 10, 1], [None, None, 5, 5]]
[[2], [10, 10], [None, 10, 1], [None, None, 5, 5]]
[[2], [10, 10], [3, 10, 1], [None, 11, 5, 5]]
15
```

<!-- grade_id: cell-f7a5d6c26e95dab2, points: 1, grade: true, solution: false -->
```python
import numpy as np
assert(bottom_up_coin_game([2, 10, 1, 5]) == 15)  
assert(bottom_up_coin_game([2, 10, 1, 5, 5]) is None)  
assert(bottom_up_coin_game([]) == 0)  
```

Output:
```
[[2], [None, 10], [None, None, 1], [None, None, None, 5]]
[[2], [10, 10], [None, 10, 1], [None, None, 5, 5]]
[[2], [10, 10], [None, 10, 1], [None, None, 5, 5]]
<ipython-input-61-b7e49a2da0fa> in <module>
      1 import numpy as np
----> 2 assert(bottom_up_coin_game([2, 10, 1, 5]) == 15)
      3 assert(bottom_up_coin_game([2, 10, 1, 5, 5]) is None)
      4 assert(bottom_up_coin_game([]) == 0)

<ipython-input-57-711796946879> in bottom_up_coin_game(A)
     37         if i >1:
     38             for j in range(i,n):
---> 39                 sol[j][-(i+1)]=max(A[j]+min(sol[j][-i],sol[j-1][-(i+1)]),A[j-i]+min(sol[j][-i],sol[j-1][-(i+1)]))
     40 
     41             #         for i in range(2,n):

IndexError: list index out of range
```

## [Optional]  Question 2. 

Complete the function `print_strategy` to print out a solution to the coin game. Completing the helper function `extended_bottom_up_coin_game` to use in `print_strategy` is optional.

<!-- grade_id: cell-4583755bc245e1e9, points: 0, grade: false, solution: true -->
```python
def print_strategy(A):
    """
    Print coins to pick for the first player
    
    Inputs:
    - A: list of floats, values of the coins, of even length.
    
    Outputs:
    - max_val: float, maximum possible amount of money that can be won by the first 
    player. 
        * max_val is None if the length of list A is odd.
        * max_val is 0 if list A is empty (no coins)
    - out: list of values of the coins that the first player picks in that
    order.
        * out is None if max_val is 0 or None.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-75ce551c5185e374, points: 0, grade: true, solution: false -->
```python
assert(print_strategy([4,6,5,20,4,6,4,8])[0] == 40)
assert(print_strategy([4,6,5,20,4,6,4,8])[1] == [8, 6, 6, 20])
```

```python

```

```python
#Korede's solution

def bottom_up_coin_game(A):
    """
    Returns the maximum possible amount of money the first-move player can win,
    given an array of coin values. The function solves this using a bottom-up 
    approach.
    
    Inputs:
    - A: list of floats, values of the coins, of even length. 
    
    Outputs:
    - max_val: float, maximum possible amount of money that can be won by the first 
    player. max_val is None when the length of list A is odd.
    """
    
    #Determine the number of coins
    n = len(A)
    
    #If n is an odd number, return None
    if n%2 == 1: 
        return None
    
    #If there are no coins, return 0 since no money can be made
    elif n==0:
        return 0
    
    #Initialize a matrix to store the solutions to subproblems
    #Rows are i and columns are j (referring to the prompt)
    solution_matrix =[[None for i in range(n)] for j in range(n)] 
    
    for length in range(n): #Iterating through the different lengths of the sublists/ size of subproblems
        for i in range(n - length):
                    
            j= i + length
            
            #If the just one coin, store the value since that's the only choice
            if length== 0:
                solution_matrix[i][j] = A[i]
                
            #If two coins, store the larger value
            elif length == 1:
                solution_matrix[i][j] = max(A[i],A[i+1])
                
            #If more than two coins
            else:
                solution_matrix[i][j] = max(A[i]+min(solution_matrix[i+2][j],solution_matrix[i+1][j-1]),
                                           A[j]+min(solution_matrix[i][j-2],solution_matrix[i+1][j-1]))
                
    return solution_matrix[0][-1]
```